---
title_meta  : Chapter 1
title       : Insert the chapter title here
description : What is this chapter about? Add here your description

---
## Create Courses for DataCamp

This is the assignment. Here, you can explain statistical concepts or R functionality. Assignments are written in (R) Markdown.

Use [Markdown](http://en.wikipedia.org/wiki/Markdown) for formatting your content.

You can do the following things with Markdown syntax:

- Make a list (as you are doing now)
- Emphasize important words: **bold** and *italics*.
- Add inline r-code with backticks: `x <- 4`. This code will not be evaluated.
- Incorporate inline LaTeX style formulas: $1 + 1 = 2$
- Link to interesting content: [this is a link](http://en.wikipedia.org/wiki/Markdown).

If you have several lines of R code you want to include in the assignment, use a code block:

```
x <- 5
x <- 7
```

If you want to include a block LaTeX style formula, use double dollar signs:

$$\frac{E(R_i) - R_f}{\beta_i} = E(R_m) - R_f$$

This exercise is included to explain the different elements of an interactive DataCamp exercise. The next exercise will give an actual example of an interactive DataCamp exercise.

*** =instructions
- Provide one instruction to your students/employees/... on what specific problem they should solve in this exercise
- Provide a second instruction here.
- etc.

*** =hint
Hints help the students to find a correct solution for the exercise. Hints are optional, but their usage is highly recommended.

*** =pre_exercise_code
```{r,eval=FALSE}
# Pre-exercise code is executed before the exercise starts (what's in a name...)
# It can be used to:
# Pre-load a dataset
# Render graphs that are shown at the start of the exercise
# Initialize other variables
# ...
```

*** =sample_code
```{r,eval=FALSE}
# The sample code is what is shown in the students editor by default
```

*** =solution
```{r,eval=FALSE}
# The solution code solves the assignment. Make sure the solution passes the submission correctness test!
```

*** =sct
```{r,eval=FALSE}
# The Submission Correctness Test code
# It takes as input the students workspace, his code and what was printed to the console during the exercise.
# Based on these inputs the SCT returns:
# (i) whether the students code was correct
# (ii) a message to the student to either help him find what was wrong or congratulate him on finding the solution.
success_msg("Well done! Continue to see some more examples!")
```

---
## Variable assignment

This simple exercise introduces you to variable assignment.
For variable assignment, use `<-`. For example:

```
a <- 3
b <- a + 2
```

*** =instructions
- Assign the value 5 to the variable `x`.
- We have already loaded a variable `y`, which is a data frame, in the workspace for you. Now, assign this variable `y` to a new variable `z`

*** =hint
Just assign 5 to the variable `x`; use the `<-` operator.

*** =pre_exercise_code
```{r,eval=FALSE}
y <- data.frame(a = c(1,2,3), b = c("this","is","me"))
```

*** =sample_code
```{r,eval=FALSE}
# Assign the value 5 to 'x'


# Assign 'y' to 'z'

```

*** =solution
```{r,eval=FALSE}
# Assign the value 5 to 'x'
x <- 5

# Assign 'y' to 'z'
z <- y
```

*** =sct
```{r,eval=FALSE}
test_object("x", undefined_msg = "Are you sure you have defined the variable <code>x</code>?",
            incorrect_msg = "Make sure to correctly specify the variable <code>x</code>")
test_object("z", undefined_msg = "Are you sure you have defined the variable <code>y</code>?",
            incorrect_msg = "Make sure to correctly specify the variable <code>y</code>")
success_msg("Awesome job! Continue to the next exercise")
```

--- type:MultipleChoiceExercise
## This is a multiple choice exercise!

The multiple choice exercise (remark the `type:` specification after the triple dashes above the title) does not need sample code, nor a solution.

*** =instructions
- Wrong option
- Correct option 
- Wrong option
  
*** =hint
Hint here. 

*** =pre_exercise_code
```{r,eval=FALSE}
```

*** =sct
```{r,eval=FALSE}
test_mc(correct = 2, feedback_msgs = c("Wrong, try again.","Good job!","Try again please"))
```

--- type:VideoExercise
## This exercise contains a video!

You can add a description here, but this is optional.

*** =video_link
```{r,eval=FALSE}
//player.vimeo.com/video/108225030
```

--- type:MarkdownExercise
## Markdown Exercise Title
 
The assignment of this markdown exercise (remark the `type:` specification after the triple dashes above the title) comes here.
Custom CSS is an easy way to add branding to your reports. 

*** =instructions
- The `styles.css` file to the right contains some example CSS that will change the appearance of your report. Modify the header of the .Rmd report to use the CSS, and then render the report.
- In the code chunk, assign `5` to the variable `x`.

*** =hint 
This exercise only expect you to edit the YAML header, not the actual .css file.

*** =pre_exercise_code
```{r eval=FALSE}
```

*** =sample_code
```{r eval=FALSE}
{{{my_document.Rmd}}}
---
title: "Tryout"
output: html_document
---

This sample R Markdown document contains a list:
  
* A
* B

It also contains some code:
  
_tbt_{r}
x <-
_tbt_
```

```{r, eval = FALSE}
{{{styles.css}}}
h1{
  color: white;
  padding: 10px;
  background-color: #3399ff
}

ul {
  list-style-type: square;  
}
```

*** =solution

```{r eval=FALSE}
{{{my_solution.Rmd}}}
---
title: "Tryout"
output: 
  html_document:
    css: styles.css
---

This sample R Markdown document contains a list:
  
* A
* B

It also contains some code:
  
_tbt_{r}
x <- 5
_tbt_
```

```{r, eval = FALSE}
{{{styles.css}}}
h1{
  color: white;
  padding: 10px;
  background-color: #3399ff
}

ul {
  list-style-type: square;  
}
```

*** =sct
```{r eval=FALSE}
test_error()
test_rmd_file({
  test_yaml_header(options = "title")
  test_yaml_header(options = "output.html_document.css", 
                   not_called_msg = "Have you specified the <code>css</code> option under <code>html_document</code>, which in turn is under option <code>output</code>?", 
                   incorrect_msg = "Make sure to set the <code>css</code> option correctly.")
  test_rmd_group(2, {
    test_object("x", 
                undefined_msg = "Have you defined the variable <code>x</code> inside the code chunk?", 
                incorrect_msg = "Make sure to define the variable <code>x</code> correctly inside the code chunk.")
  })
})
success_msg("Nice work! It's easy to add styling to your report!")  
```
